# Lab 6: Examining the Therapeutic Touch
This assignment involves using statistical modeling, hypothesis testing, and simulation to analyze data and draw conclusions.

**Attention:**

1. The noteook:
    1. Use **Jupyter Notebook** as instructed. We do not have the resources to support <font color="darkred">**Anaconda** or **VS Code**</font>. 
    2. Provide your answers in the designated spaces.
    3. Do not reassign variables in the notebooks! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on (e.g., as an intermediate variable). Otherwise, you will fail the tests that you thought you were passing previously!
2. Grading:
    1. Points will be scaled in Canvas.
    2. Most "tests" in this assignment test the **format and data types** of your answers, not their **correctness**. Passing **100%** of the **tests** does not mean your **grade** will be 100%.
3. Academic honor:
    1. Academic honor is important: **DO NOT CHEAT!** 
    2. <font color='darkred'>**Use AI to LEARN, not to cheat**. **DO NOT copy-n-paste**</font>: You will be competing with people who are as good as you in AI, and you will have the competitive edge when you are one of the few who know how to do things without AI. 
    3. Collaboration and using resources are encouraged, but DO NOT directly share answers. Upon typing up answers, do it yourself. 
4. Getting help:
    1. Come to the TA's help sessions and the instructor's office hours for help and clarification.
    2. Text a screenshot of your technical issue to the instructor.
    3. Emailing the instructor and the TA. (Emailing is too slow for technical problem-solving, so do not expect quick responses.)
5. <font color="darkred">**Time management**</font>: Start early in the assignment cycle and allot enough time for it.

**Reference Materials**:
- The [Python Reference sp25](https://www.data8.org/sp25/reference/) or [Python Reference sp24](https://www.data8.org/sp24/reference/).
- The [Data8 datascience Reference](https://www.data8.org/datascience/tables.html) is very helpful with **syntax and examples**. For example, the [Table Functions and Methods](https://www.data8.org/datascience/reference-nb/datascience-reference.html) and [Tables](https://www.data8.org/datascience/tables.html)

**Recommended Readings**: 
* [11.1. Assessing a Model](https://introdsm.org/chapters/11/1/Assessing_a_Model.html)
* [10.3. Empirical Distribution of a Statistic](https://introdsm.org/chapters/10/3/Empirical_Distribution_of_a_Statistic.html)
* [10.4. Random Sampling in Python](https://introdsm.org/chapters/10/4/Random_Sampling_in_Python.html)


```python
# Initialize Otter
import otter
grader = otter.Notebook("lab06.ipynb")
```

```python
# Run this cell, but please don't change it.
 
# These lines import the Numpy and Datascience modules.
import numpy as np
from datascience import *

# These lines do some fancy plotting magic
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
from matplotlib import patches
from ipywidgets import interact, interactive, fixed
import ipywidgets as widgets

```

## 1. What is the Therapeutic Touch

The Therapeutic Touch (TT) is the idea that everyone can feel the Human Energy Field (HEF) around individuals.  Those who practice TT have described different people's HEFs as "warm as Jell-O" and "tactile as taffy." 

TT was a popular technique used throughout the 20th century that was toted as a great way to bring balance to a person's health. Certain practitioners claim they have the ability to feel the HEF and can massage it in order to promote health and relaxation in individuals.

### Emily Rosa

[Emily Rosa](https://en.wikipedia.org/wiki/Emily_Rosa) was a 4th grade student who was very familiar with the world of TT, thanks to her parents, who were both medical practitioners and skeptics of TT.

For her 4th grade science fair project, Emily decided to test whether or not TT practitioners could truly interact with a person's HEF. She later went on to publish her work in TT, becoming the youngest person to have a research paper published in a peer reviewed medical journal.

### Emily's Experiment

Emily's experiment was clean, simple, and effective. Due to her parents' occupations in the medical field, she had wide access to people who claimed to be TT practitioners. 

Emily took 21 TT practitioners and used them for her science experiment. She would take a TT practitioner and ask them to extend their hands through a screen (which they can't see through). Emily would be on the other side and would flip a fair coin. Depending on how the coin landed, she would put out either her left hand or her right hand. The TT practitioner would then have to answer which hand Emily put out. If a pracitioner could truly interact with a person's HEF, it would be expected that they answered correctly.

Overall, through 210 samples, the practitioner picked the correct hand 44% of the time. 

Emily's main goal here was to test whether or not the TT practicioners' guesses were random, like the flip of a coin. In most medical experiments, this is the norm. **We want to test whether or not the treatment has an effect, *not* whether or not the treatment actually works.**

We will now begin to formulate this experiment in terms of the terminology we learned in this course. 

```python
### The study tested the ability of 21 TT practitioners to detect the HEF when they 
### were not looking. Rosa asked each of the practitioners to sit at a table and 
### extend their hands through a screen. On the other side of the screen, Rosa 
### randomly selected one of the TT practitioners' hands and held her own hand 
### over it. The TT practitioners were then asked which of their hands detected 
### Rosa's HEF. Subjects were each given ten tries, but they correctly located 
### Rosa's hand an average of only 4.4 times.
### https://en.wikipedia.org/wiki/Emily_Rosa
```

**Question 1.1**: Describe Emily’s [model](https://introdsm.org/chapters/11/1/Assessing_a_Model.html) for how likely the TT practitioners are to choose the correct hand. What alternative model is her model meant to discredit?

If you are able, check in with fellow peers, the discussion forum, or your lab TA/AIs, to come to a conclusion.


_Type your answer here, replacing this text._

**Question 1.2:** Remember that the practitioner got the correct answer 44% (0.44) of the time. According to Emily's model, on average, what proportion of times do we expect the practitioner to guess the correct hand? Make sure your answer is a number between 0 and 1. 


```python
expected_proportion_correct = .5
expected_proportion_correct
```

Output:
```
0.5
```

```python
grader.check("q1_2")
```

Output:
```
q1_2 results: All test cases passed!
```

The goal now is to see if our deviation from this expected proportion of correct answers is due to something other than chance. 

**Question 1.3:** We usually use a statistic to help determine which model the evidence points towards. What is a statistic that we can use to compare outcomes under Emily’s model to what was observed? Assign `valid_stat` to an array of integer(s) representing test statistics that Emily can use: 

1. The difference between the expected percent correct and the actual percent correct
2. The absolute difference between the expected percent correct and the actual percent correct
3. The sum of the expected percent correct and the actual percent correct

**NOTE:** Make sure to use `make_array` to create your array of integer(s)!

> *Hint*: What should the domain (possible x values) be for the distribution of our test statistics?


```python
valid_stat = make_array( 1, 1)
valid_stat
```

Output:
```
array([1, 1])
```

```python
grader.check("q1_3")
```

Output:
```
q1_3 results: All test cases passed!
```

**Question 1.4:** Why is the statistic from Question 1.3 the appropriate choice for comparing outcomes in Emily's experiment? How does it relate to the models you defined in Question 1.1?


_Type your answer here, replacing this text._

**Question 1.5:** Define the function `statistic` which takes in an expected proportion and an actual proportion, and returns the value of the statistic chosen in Question 1.3. Assume that the argument takes in proportions, but  return your answer as a percentage. 

*Hint:* Remember we are asking for a **percentage**, not a proportion. 


```python
def statistic(expected_prop, actual_prop):
    ...
```

```python
grader.check("q1_5")
```

**Question 1.6:** Use your newly defined function to calculate the observed statistic from Emily's experiment. 


```python
observed_statistic = ...
observed_statistic
```

```python
grader.check("q1_6")
```

**Is this observed statistic consistent with what we expect to see under Emily’s model?**

In order to answer this question, we must simulate the experiment as though Emily's model was correct, and calculate our statistic for every simulation.

### `sample_proportions`

`sample_proportions` can be used to randomly sample from multiple categories when you know the proportion of data points that are expected to fall in each category. `sample_proportions` takes two arguments: the sample size and an array of proportions corresponding to each category in the population (should sum to 1).

Consider flipping a fair coin, where the two outcomes (coin lands heads and coin lands tails) occur with an equal chance. We expect that half of all coin flips will land heads, and half of all coin flips will land tails.

Run the following cell to see the simulation of 10 flips of a fair coin. Let the first item of `coin_proportions` be the proportion of heads and the second item of `coin_proportions` be the proportion of tails.

*Observe what happens when you run this cell multiple times—the proportion of coin flips that land heads and tails appears to change, as you are simulating flipping 10 coins each time!*

```python
coin_proportions = make_array(0.5, 0.5) 
ten_flips = sample_proportions(10, coin_proportions)
ten_flips
```

`sample_proportions` returns an array that is the same length as the proportion array that is passed through. It contains the proportion of each category that appears in the sample. 

In our example, the first item of `ten_flips` is the simulated proportion of heads and the second item of `ten_flips` is the simulated proportion of tails.

```python
simulated_proportion_heads = ten_flips.item(0)
simulated_proportion_tails = ten_flips.item(1)

print("In our simulation, " + str(simulated_proportion_heads) + " of flips were heads and " \
      + str(simulated_proportion_tails) + " of flips were tails.")
```

**Question 1.7:** To begin simulating, we should start by creating a representation of Emily's model to use for our simulation. This will be an array with two items in it. The first item should be the proportion of times a TT practictioner picks the correct hand, assuming that Emily’s model was correct. The second item should be the proportion of times, under the same assumption, that the TT practitioner picks the incorrect hand. Assign `model_proportions` to this array. 

After this, we can simulate 210 hand choices, as Emily evaluated in real life, and find a single statistic to summarize this instance of the simulation. Use the `sample_proportions` function and assign the **proportion of correct hand choices** (out of 210) to `simulation_proportion_correct`. Lastly, use your `statistic` function to assign `one_statistic`  to the value of the statistic for this one simulation.

*Hint:* `sample_proportions` usage can be found on the [Python Reference](https://www.data8.org/sp24/reference/).


```python
# This saves the random state of our code so that we can 
# generate the same numbers each time we run the code.
# Please do not change this line. 
np.random.seed(16)

model_proportions = ...
simulation_proportion_correct = ...
one_statistic = ...
one_statistic
```

```python
grader.check("q1_7")
```

**Question 1.8:** Let's now see what the distribution of statistics is actually like under Emily's model. 

Define the function `simulation_and_statistic` to take in the `model_proportions` array and the expected proportion of times a TT practitioner would guess a hand correctly under Emily's model. The function should simulate Emily running through the experiment 210 times and return the statistic of this one simulation. 

*Hint:* This should follow the same pattern as the code you did in the previous problem.  

```python
def simulation_and_statistic(model_proportions, expected_proportion_correct):
    '''Simulates 210 TT hand choices under Emily’s model. 
    Returns one statistic from the simulation.'''
    ...
```

```python
num_repetitions = 1000

simulated_statistics = ...

for ... in ...:
    ...
```

```python
grader.check("q1_8")
```

Let's view the distribution of the simulated statistics under Emily's model, and visually compare where the observed statistic lies relative to the simulated statistics.

```python
t = Table().with_column('Simulated Statistics', simulated_statistics)
t.hist()
plt.scatter(observed_statistic, 0, color='red', s=100, zorder=2);
```

We can make a visual argument as to whether we believe the observed statistic is consistent with Emily’s model. Here, since larger values of the test statistic suggest the alternative model (where the chance of guessing the correct hand is something other than 50%), we can formalize our analysis by finding what proportion of simulated statistics were as large or larger than our observed test statistic (the area at or to the right of the observed test statistic). If this area is small enough, we’ll declare that the observed data are inconsistent with our simulated model. Here is the [link](https://inferentialthinking.com/chapters/11/1/Assessing_a_Model.html) to the section in the textbook.

**Question 1.9:** Calculate the proportion of simulated statistics in Question 1.8 greater than or equal to the observed statistic. 

*Hint:* `np.count_nonzero` usage can be found [here](https://www.data8.org/sp24/reference/#array-functions-and-methods).


```python
proportion_greater_or_equal = ...
proportion_greater_or_equal
```

```python
grader.check("q1_9")
```

By convention, we often compare the proportion we just calculated to 0.05. If the proportion of simulated statistics greater than or equal to the observed statistic is sufficiently small (less than or equal to 0.05), then this is evidence against Emily's model. Conceptually, you may think of this as the case where less than 5% of simulated values are as far or farther away from what we had expected. If this is not the case, we don’t have any reason to doubt Emily’s model. 

This should help you make your own conclusions about Emily Rosa's experiment. 

Therapeutic touch fell out of use after this experiment, which was eventually accepted into one of the premier medical journals. TT practitioners hit back and accused Emily and her family of tampering with the results, while some claimed that Emily's bad spiritual mood towards therapeutic touch made it difficult to read her HEF. Whatever it may be, Emily's experiment is a classic example about how anyone, with the right resources, can test anything they want!

**Question 1.10:** Now, take some time to reflect on the quetions below and then, discuss with your peers or take a look at the discussions on the Ed post for this lab.

1. Is the data more consistent with Emily' model (practioners were randomly guessing)?
2. What does this mean in terms of Emily's experiment? Do the TT practitioners' answers follow an even chance model or is there something else at play? 

Did you talk to your peers or look at the discussion forum? (True/False)


```python
peer_talk = ...
peer_talk
```

```python
grader.check("q1_10")
```

<p>&nbsp;</p>

`Good Job! You are done with this assignment!`

## Submission

After you have completed the assignment, do the following to submit it:

1. **Save** the notebook file (File ==> Save Notebook) (or just the Save icon)
2. Go to the notebook menu, choose Kernel ==> **Restart Kernel and Run All Cells**.
3. Scrolling around to make sure everything works fine without unexpected error messages. 
4. **Save** the notebook file (File ==> Save Notebook, or use the Save icon).
5. Use the Jupyter Notebook dashboard/Homepage to create a **duplicate** of this notebook and then **rename** it from *assignment*.ipynb_**copy** (e.g., a01.ipynb_copy) to *assignment_**FIRSTNAME_LASTNAME***.ipynb (e.g., a01_TSANGYAO_CHEN.ipynb) to be graded. That way you will be able to keep your original file.
6. **Upload** your <font color="blue">*assignment_**FIRSTNAME_LASTNAME***.ipynb</font> to Canvas.

Tests for q1_10 (points: None)
```
>>> peer_talk == True
True
```


Tests for q1_2 (points: None)
```
>>> 0 <= expected_proportion_correct <= 1
True
```
```
>>> expected_proportion_correct == 0.5
True
```


Tests for q1_3 (points: None)
```
>>> type(valid_stat) == np.ndarray
True
```
```
>>> int(sum(valid_stat))
2
```


Tests for q1_5 (points: None)
```
>>> int(round(statistic(.5,.5) + statistic(.4,.1),1))
30
```
```
>>> int(statistic(.4,.1) - statistic(.1,.4))
0
```


Tests for q1_6 (points: None)
```
>>> int(round(observed_statistic,2))
6
```


Tests for q1_7 (points: None)
```
>>> len(model_proportions) % 2 == 0
True
```
```
>>> len(np.unique(model_proportions))
1
```
```
>>> sum(model_proportions) == 1
True
```
```
>>> type(simulation_proportion_correct) == float
True
```
```
>>> round(simulation_proportion_correct, 2)
0.49
```
```
>>> round(one_statistic, 2) - 0.95 < 0.05
True
```


Tests for q1_8 (points: None)
```
>>> len(simulated_statistics) == 1000
True
```
```
>>> np.all(simulated_statistics <= 30)
True
```
```
>>> np.all(simulated_statistics >= 0)
True
```
```
>>> 0 <= simulation_and_statistic(model_proportions, expected_proportion_correct) <= 25
True
```


Tests for q1_9 (points: None)
```
>>> 0 <= proportion_greater_or_equal <= 1
True
```
```
>>> proportion_greater_or_equal*1000 == np.count_nonzero(simulated_statistics >= observed_statistic)
True
```
