# Lab 8: Estimation

**Attention:**

1. **Using the Notebook**
    - Use Jupyter Notebook as instructed. We do **not** have the resources to support Anaconda or VS Code. 
    - Provide your answers in the designated spaces.
    - Do not reassign variables! For example, if you use `max_temperature` for one question, do not reassign it later on (e.g., as an intermediate variable). Doing so may cauae you to fail tests you passed previously!
2. **Grading**:
    - Points may be scaled in Canvas.
    - Most tests check the format and data types of your answers, not their correctness. Passing the tests does not guarantee the grade will be 100%.
3. Academic honor:
    - **DO NOT CHEAT!** Academic honor matters.  
    - <font color='darkred'>**Use AI to LEARN, not to cheat**. **DO NOT copy-and-paste**</font>. Your advantage comes from your knowledge and skills on the subject. Do not rely on AI. 
    - Collaboration and using resources are encouraged, but **DO NOT directly share answers**. Always answer the questions and solve the problems by yourself. 
4. Getting help:
    - Attend the TA's help sessions and the instructor's office hours.
5. <font color="darkred">**Time management**</font>: Start early in the assignment cycle and allocate enough time to complete the assignment properly.

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab08.ipynb")
```

```python
# Run this cell, but please don't change it.
 
# These lines import the Numpy and Datascience modules.
import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
# import matplotlib
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore')
```

# 1. Percentiles

Numerical data (array, list, sequence, collection, etc) can be sorted in increasing or decreasing order. Thus the values of a numerical data set have a **rank order**. The value at a particular rank is called *p*th "percentile" of the data.

#### 1.1. Calculating Percentiles

First, let's create an array using the make_array() function to help us learn about percentiles. In the space below, create an numpy array with elements of **12, 17, 6, 9, 7** and call it **nums**.

**q1.1.1 (0.5)** 

create your array here and call it nums


```python
...
```

**q1.1.2 (0.5)**

- jupyter notebook evaluate expressions interactively just like the Python shell
- type your array name below and run the cell to see the array

```python
...
```

**q1.1.3 (0.5)**
- it is important to always know the data type of your data
- show the type of nums below by using the type function

```python
...
```

```python
### q1.1.4 (0.5)
### sort the array in increasing order using the numpy sort function and
### reassign the ordered array to the same variable name nums
...
```

**q1.1.5 (2 points)**

the 50th percentile of a sequence is its median.
your favorite isntructor has created a one-liner to find the median of a sequence.
the one-liner actually almost works if he has not made 3 mistakes. correct your instructor by:
1. fix the code.
2. replace the text below with your descriptive corrections in the format of change the XXX to YYY because of ZZZ. 

```python
# nums[float(np.floor(50/100 * len(nums)))]
...            ### fixed code in this line
```

<font color='blue'>_correction 1 is about data type. replace this text with your correction._</font>

<font color='green'>_correction 2 is about the *p*th percentile being the smallest value in the collection that is at least as learge as p% of the collection._</font>

<font color='red'>_correction 3 is about subscripting to reference the right element._</font>


**q1.1.6 (1 point)**
- eventually, we use the percentile function to find our median.
- use the percentile function to find the median of the array numbers with elements of 1, 2, 3, 4, 4, 4, 4, 4, 5, 6, 7 

```python
### create the array using the make_array function and assign to variable name numbers
### get the median of the list using the percentile function
numbers = ...
median_of_numbers = ...          ### get the median 
```

#### 1.2. Visulization

**q1.2.1** 

set your path_data variable correctly to access the datasets in your data folder. 


```python
path_data = ...
```

**q1.2.2**

load the dataset (the "scores_and_sections" table in file 'scores_by_section.csv')


```python
scores_and_sections = ...

```

**q1.2.3** 

you may want to take a look at the table before moving on: use the show() function to show 5 rows


```python
...          ### show 5 rows of the table
```

**q1.2.4**

find the median of the Midterm scores using the percentile() function


```python
the_median = ...
```

**q1.2.5** 

- create a histogram of the midterm scores
- use np.arange(-0.5, 25.6, 1) for bins

```python
scores_and_sections.hist...

### plot the_median using following code
parameter_green = '#32CD32'
plots.ylim(-0.002, .08)
plots.scatter(the_median, 0, color=parameter_green, s=40, zorder=2)
```

Note: Your histogram should look like this:

<img src="midterm.jpg" width = 350>

What function puts the green dot on the histogram? 

<font color='blue'> ... _replace this text with your answer._</font>

## 2. The Bootstrap Process

We will use bootstrap to estimate the population median. The idea is, we have some delay data in minutes about United Airline. We want to estimate how many minutes delay we will experence if we take United Airline. (This can be interpreted as we are treating the United Airline delay dataset as the population dataset, which is usually not available.)

**q2.1.1** 

load the dataset 'united_summer2015.csv' and call it united

```python
united = ...
```

**q2.1.2** 

show the first three lines of sf2019 using the show() function

```python
...
```

**q2.1.3** stats()

The stats() function gives us basic descriptive statistics about the table.

```python
### just run this cell
united.stats()
```

**q2.1.4** 

Based on the output of the code, above, which of the following would be a good bins for displaying a histogram of the Delay minutes? Place your answer in the cell below. 

1. np.arange(-20, 600, 10)
2. np.arange(-20, 580, 10)
3. np.arange(-10, 580, 10)
4. np.arange(-10, 600, 10)

<font color='blue'> ... _replace this text with your answer._</font>

**q2.2** Visualization

Use the bins you choose above to create a hitogram to show the "Delay" information in the united table.

```python
...         ### use your bin
...         ### create histogram
```

**q2.3.** Data cleasing

Obvriously the right tail of this distribution is very long and we should consider reasonable data cleasing. Supposed you are a domain expert and you think we should use 150 minutes as a cutoff point for this task. 

Use a line of code to remove the rows with greator than 150 minutes in delay. Reassign the new table to  name it united_cleaned.

```python
### write your line of code here
united_cleaned = ...
```

```python
### just run this cell to see the stats
united_cleaned.stats()
```

**q2.3.3** 

Use the num_row attribute to check the number of rows in the united_cleaned table.

```python
### type your code here
...
```

We want to visualize the cleaned data

```python
### NOTE THAT we defined a new bins variable here.
### just run the code
new_bins = np.arange(-20, 150, 5)
united_cleaned.select('Delay').hist(bins=new_bins)
```

**q2.4.** The paramter: median

Calculate the median of the minutes of delay using one line of code and call it pop_median.

```python
pop_median = ...
pop_median
```

**q2.4.** Random sampling

- draw a sample of 500 delays at random without replacement from the cleaned table
- visualize the data to see how different it is from the population

```python
our_sample = ...    ### 500 random samples
...                   ### histogram using the new bin
```

**q2.5.** Estimate the parameter

- find the median delay of the sampled employees serve as our estimate of the parameter
- Use a one-line code to calculate the estimated median of our sample.
- the result should not be too far away from the pop_median

```python
...
```

**q2.6.** The bootstrap: A resampled median

- use a one-line code to resample the our_sample distribution and create a delay histogram.

```python
### resample our_sample and call it resample_1
...
```

- just to check how different this resampling is from the original sample (our_sample)
- create a delay histogram of resample_1 using new_bins

```python
### histogram of resample_1 on delay
...
```

**q2.7.** 

calculate the estiamte of the population median and call it resample_median_1

```python
resample_median_1 = ...    ### resample_median_1
resample_median_1                                                 ### evaluate resample_median_1
```

**q2.8** resample the sample many tims to create an empirical distution of the estimates:

- define a function, call it one_bootstrap_median, to retunr one bootstrapped median of delays based on the original random sample "our_sample". 

```python
def ...                                            ### function header
    ...                                 ### sample our_sample to a new table, call it resample_tbl
    ... ### calculate the median, call it bootstrapped_median
    return ...                                        ### return the obtained median
```

```python
### calling this function multiple times (Control+Enter) should help your understanding about resampling
one_bootstrap_median()
```

**q2.9.** bootstrap empirical distribution of the sample median

1. we are asking for 5000 repetitions. Create a variable for this and call it num_repetitions.
2. create an empty array called bstrap_medians.
3. design a for loop to call the one_bootstrap_median function 5000 times and save the medians to the array.

```python
...                ### define num_repetitions
...          ### create the empty array called bstrap_medians
...   ### a for loop to run num_repetitions times
    ...     ### append the medians to the array
```

```python
### use the stats() function for basic information
### run the code and observe the table
### we can use the min/max for our bins
resampled_medians.stats()
```

**q2.10.** turn the median array to a table so that we can visualize the distribution

- create and call the table resampled_medians
- create a histogram with bins (0.5, 5.1, 1) based on the stats above

```python
resampled_medians = ...
resampled_medians.hist(...)

# plotting the pop_median; you can ignore this code
parameter_green = '#32CD32'
plots.ylim(-0.05, 1)
plots.scatter(pop_median, 0, color=parameter_green, s=40, zorder=2);
```

**q2.11.** do the estimates capture the parameter?

- calculating the “middle 95%” interval of resampled medians: 

```python
left = ...         ### the left 2.5%
right = ...       ### the right 97.5%
make_array(left, right)
```

**q2.12** 

is the population median between the two numbers (left and right)?

_<font color='blue'>answer the question; replacing this text</font>_

**q2.13.** plotting the histogram

```python
### run the code
resampled_medians.hist(bins = np.arange(0.5, 5.1, 1))

# plotting parameters; you can ignore this code
plots.ylim(-0.05, 1)
plots.plot([left, right], [0, 0], color='yellow', lw=3, zorder=1).  
plots.scatter(pop_median, 0, color=parameter_green, s=40, zorder=2);
```

Your histogram should look like this:

<img src='medians_intervals.jpg' width=450>

- The “middle 95%” interval of estimates captured the parameter in our example. 

<p>&nbsp;</p>

`Good Job! You are done with this assignment!`

## Submission Instructions

Follow these steps to submit your work:

1. Save your notebook file:
   - Click on File -> Save Notebook, or simply click on the Save icon.
2. Restart Kernel and Run All Cells:
   - In the menu, go to Kernel → Restart Kernel and Run All Cells.
3. Verify your work:
   - Scroll through the notebook to ensure everything runs smoothly without unexpected error messages. 
   - In case of interruptions caused by execution errors, use "Run ==> Selected Cell and All Below" to continue execution. 
4. Save the notebook again to confirm all changes.
5. Create a duplicate notebook for submission:
   - In the Jupyter Notebook dashboard/homepage, create a duplicate of this notebook
   - Rename the duplicate from *assignment.ipynb_copy* to *assignment_FIRSTNAME_LASTNAME.ipynb* (e.g., *a01_TSANGYAO_CHEN.ipynb*) for submission.
   - Keep your original notebook file as a backup.
6. Upload your renamed notebook to Canvas.

Tests for q1_3 (points: None)
```
>>> round(age_mean, 3) == 27.228
True
```
```
>>> round(age_sd, 3) == 5.815
True
```


Tests for q1_4 (points: None)
```
>>> round(height_mean, 3) == 64.049
True
```
```
>>> round(height_sd, 3) == 2.525
True
```


Tests for q2_1 (points: None)
```
>>> np.random.seed(8)
>>> np.round(one_sample_mean(salaries, 'salary', 100), 3) == 76699.829
True
```


Tests for q2_4 (points: None)
```
>>> 
>>> import hashlib # This imports a hashing library for the autograder.
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(set(np.array(q2_4).astype(int))) 
'3649655f87d3a44afa82f25eced24944'
```


Tests for q2_5 (points: None)
```
>>> 
>>> import hashlib # This imports a hashing library for the autograder.
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(set(np.array(q2_5).astype(int)))
'b793dd27381f7b63cada88b749c0e023'
```


Tests for q2_6 (points: None)
```
>>> import hashlib # This imports a hashing library for the autograder.
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(int(SD_of_sample_means))
'eccbc87e4b5ce2fe28308fd9f2a7baf3'
```


Tests for q2_7 (points: None)
```
>>> import hashlib # This imports a hashing library for the autograder.
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(set(np.array(pop_vs_sample).astype(int)))
'1d919a653870c713cf8bd17b9c9c65ee'
```


Tests for q2_8 (points: None)
```
>>> 
>>> import hashlib # This imports a hashing library for the autograder.
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(int(q2_8))
'c4ca4238a0b923820dcc509a6f75849b'
```
