# Homework (hw01) Instructions

## 0. .ipynb Instructions

#### Activating virtual environment
1. Go to your **terminal** (PowerShell for Windows and Terminal.app for macOS) & move (cd) into your project directory (**workspace/introdsm**)
   1. The lab computers can be inconsistent in the virtual environment and package installation.
   2. Bring your own device (BYOD) is strongly encouraged. 
2. Activate your virtual environment:
   1. For macOS, Enter "source .venv/bin/activate" 
   2. For Windows, Enter "./.venv/Scripts/activate" 
   3. You should see the ".venv" prefix at your terminal prompt when the virtual environment is activated.
3. Launch Jupyter Notebook by entering "jupyter notebook" at the command line.
4. When you are done, deactivate the virtual environment by entering "deactivate".

#### Installing packages: otter-grader & datascience
1. Ideally, you should install packages using a terminal:
   1. In a terminal, activate your virtual machine.
   2. Enter "pip install <i>package_name</i>" to install module/package/ibrary. 
2. To install packages in Jupyter Notebook, for example:   
   1. type **```%pip install datascience```** and hit Enter key (this will take a while, like minutes)
   2. type **```%pip install otter-grader```** and Enter key. 
3. Comment out the installation line after installation is done. 

```python
# Initialize Otter
import otter
grader = otter.Notebook("hw01.ipynb")
```

# Homework 1: Causality and Expressions

Please complete this notebook by filling in the cells provided. Before you begin, run the previous cell to load the provided tests.

**Recommended Readings:**

- [What is Data Science?](http://www.inferentialthinking.com/chapters/01/what-is-data-science.html)
- [Causality and Experiments](http://www.inferentialthinking.com/chapters/02/causality-and-experiments.html) 
- [Programming in Python](http://www.inferentialthinking.com/chapters/03/programming-in-python.html)

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!

**Deadline:**

This assignment is **due Wednesday, 1/24 at 5:00pm PT**. Submissions after this time will be accepted for 24 hours and will incur a 20% penalty. Any submissions later than Thursday, 1/25 at 5:00pm PT will not be accepted unless an extension has been granted. Turn it in by Tuesday, 1/23 at 5:00pm PT for 5 extra credit points. For more information check our [policies](http://data8.org/sp24/policies/) page.

**Note: This homework has hidden tests on it. That means even though tests may say 100% passed, it doesn't mean your final grade will be 100%. We will be running more hidden tests for correctness once everyone turns in the homework.**

Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. Refer to the [policies](http://data8.org/sp24/policies/#academic-honesty) page to learn more about how to learn cooperatively.

You should start early so that you have time to get help if you're stuck. Office hours are held Monday-Friday. The schedule appears on [our website](http://data8.org/sp24/officehours/).

## 1. Scary Arithmetic

<!-- BEGIN QUESTION -->

**Question 1.1** An ad for ADT Security Systems says,

> "When you go on vacation, burglars go to work [...] According to FBI statistics, over 25% of home burglaries occur between Memorial Day to Labor Day."

Do the data in the ad support the claim that burglars are more likely to go to work during the time between Memorial Day to Labor Day? Please explain your answer. **(6 Points)**

**Note:** You can assume that "over 25%" means only slightly over. Had it been much over, say closer to 30%, then the marketers would have said so.

**Note:** Memorial Day is observed on the last Monday of May and Labor Day is observed on the first Monday of September.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 2. Characters in Little Women


In lecture, we counted the number of times that the literary characters were named in each chapter of the classic book, [*Little Women*](https://inferentialthinking.com/chapters/01/3/1/Literary_Characters.html?highlight=little%20women). In computer science, the word "character" also refers to a letter, digit, space, or punctuation mark; any single element of a text. The following code generates a scatter plot in which each dot corresponds to a chapter of *Little Women*. The horizontal position of a dot measures the number of periods in the chapter. The vertical position measures the total number of characters.

```python
# Just run this cell.

# This cell contains code that hasn't yet been covered in the course,
# but you should be able to interpret the scatter plot it generates.

from datascience import *
from urllib.request import urlopen
import numpy as np
import ssl
context = ssl._create_unverified_context()
%matplotlib inline


little_women_url = 'https://www.introdsm.org/data/little_women.txt'
chapters = urlopen(little_women_url, context=context).read().decode().split('CHAPTER ')[1:]
text = Table().with_column('Chapters', chapters)
Table().with_columns(
    'Periods',    np.char.count(chapters, '.'),
    'Characters', text.apply(len, 0)
    ).scatter(0)
```

**Question 1.** Around how many periods are there in the chapter with the most characters? Assign either 1, 2, 3, 4, or 5 to the name `characters_q1` below. **(4 Points)**

1. 250
2. 390
3. 440
4. 32,000
5. 40,000


**Note:** If you run into a `NameError: name 'grader' is not defined` error in the autograder cell below (and in any assignment), please re-run the first cell at the very top of this notebook!


```python
characters_q1 = ...
```

```python
grader.check("q2_1")
```

The test above checks that your answers are in the correct format. **This test does not check that you answered correctly**, only that you assigned a number successfully in each multiple-choice answer cell.

**Question 2.** Which of the following chapters has the most characters per period? Assign either 1, 2, or 3 to the name `characters_q2` below. **(4 Points)**

1. The chapter with about 60 periods
2. The chapter with about 350 periods
3. The chapter with about 440 periods


```python
characters_q2 = ...
```

```python
grader.check("q2_2")
```

Again, the test above checks that your answers are in the correct format, but not that you have answered correctly.

To discover more interesting facts from this plot, check out [Section 1.3.2](https://inferentialthinking.com/chapters/01/3/2/Another_Kind_Of_Character.html) in the textbook.

## 3. Names and Assignment Statements

**Question 1.** When you run the following cell, Python produces a cryptic error message.

```python
4 = 2 + 2
```

Choose the best explanation of what's wrong with the code, and then assign 1, 2, 3, or 4 to `names_q1` below to indicate your answer. **(4 Points)**

1. Python is smart and already knows `4 = 2 + 2`.

2. In Python, it's a rule that the `=` sign must have a variable name to its left, and `4` isn't a variable name.

3. It should be `2 + 2 = 4`.

4. I don't get an error message. This is a trick question.


```python
names_q1 = ...
```

```python
grader.check("q3_1")
```

**Question 2.** When you run the following cell, Python will produce another cryptic error message.

```python
two = 3
six = two plus two
```

Choose the best explanation of what's wrong with the code and assign 1, 2, 3, or 4 to `names_q2` below to indicate your answer. **(4 Points)**

1. The `plus` operation only applies to numbers, not the word "two".

2. The name "two" cannot be assigned to the number 3.

3. Two plus two is four, not six.

4. The name `plus` isn't a built-in operator; instead, addition uses `+`.


```python
names_q2 = ...
```

```python
grader.check("q3_2")
```

**Question 3.** Run the following cell.

```python
x = 2
y = 3 * x
x = 4
```

What is `y` after running this cell, and why? Choose the best explanation and assign 1, 2, 3, or 4 to `names_q3` below to indicate your answer. **(4 Points)**

1. `y` is equal to 6, because the second `x = 4` has no effect since `x` was already defined.

2. `y` is equal to 6, because `x` was 2 when `y` was assigned, and 3 * 2 is 6.

3. `y` is equal to 12, because `x` is 4 and 3 * 4 is 12.

4. `y` is equal to 12, because assigning `x` to 4 will update `y` to 12 since `y` was defined in terms of `x`.


```python
names_q3 = ...
```

```python
grader.check("q3_3")
```

## 4. Differences Between Majors

Berkeley’s Office of Planning and Analysis (OPA) provides data on numerous aspects of the campus. Adapted from the OPA website, the table below displays the number of degree recipients in three majors in the 2008-2009 and 2017-2018 academic years.

| Major                              | 2008-2009    | 2017-2018   |
|------------------------------------|--------------|-------------|
| Gender and Women's Studies         |      17      |    28       |
| Linguistics                        |      49      |    67       |
| Rhetoric                           |      113     |    56       |



**Question 1.** Suppose you want to find the **biggest** absolute difference between the number of degree recipients in the two years, among the three majors.

In the cell below, compute this value and call it `biggest_change`. Use a single expression (a single line of code) to compute the answer. Let Python perform all the arithmetic (like subtracting 49 from 67) rather than simplifying the expression yourself. The built-in `abs` function takes a numerical input and returns the absolute value. The built-in `max` function can take in 3 arguments and returns the maximum of the three numbers. **(5 Points)**


```python
biggest_change = ...
biggest_change
```

```python
grader.check("q4_1")
```

**Question 2.** Which of the three majors had the **smallest** absolute difference? Assign `smallest_change_major` to 1, 2, or 3 where each number corresponds to the following major:

1. Gender and Women's Studies  
2. Linguistics  
3. Rhetoric

Choose the number that corresponds to the major with the smallest absolute difference. **(4 Points)** 

_Hint:_ You should be able to answer by rough mental arithmetic, without having to calculate the exact value for each major.


```python
smallest_change_major = ...
smallest_change_major
```

```python
grader.check("q4_2")
```

**Question 3.**  For each major, define the “relative change” to be the following: $\large{\frac{\text{absolute difference}}{\text{value in 2008-2009}} * 100}$ 

Fill in the code below such that `gws_relative_change`, `linguistics_relative_change` and `rhetoric_relative_change` are assigned to the relative changes for their respective majors. **(5 Points)**


```python
gws_relative_change = (abs(...) / 17) * 100
linguistics_relative_change = ...
rhetoric_relative_change = ...
gws_relative_change, linguistics_relative_change, rhetoric_relative_change
```

```python
grader.check("q4_3")
```

**Question 4.** Assign `biggest_rel_change_major` to 1, 2, or 3 where each number corresponds to to the following: 

1. Gender and Women's Studies  
2. Linguistics  
3. Rhetoric

Choose the number that corresponds to the major with the biggest relative change. **(4 Points)**


```python
biggest_rel_change_major = ...
biggest_rel_change_major
```

```python
grader.check("q4_4")
```

## Submission

After you have completed the assignment, do the following to submit it:

1. **Save** the notebook file (File ==> Save Notebook) (or the Save icon)
2. Go to the notebook menu, choose Kernel ==> **Restart Kernel and Run All Cells**.
3. Make sure everything is fine. 
4. **Save** the notebook file (File ==> Save Notebook)
5. Use the Juypter Notebook dashboard to create a **duplicate** of the file and then **rename** it from hw01.ipynb_copy to **hw01_FIRSTNAME_LASTNAME**.ipynb to be graded. That way you will be able to keep your original file. 
6. **Upload** the file to Canvas.

Tests for q2_1 (points: [0])
```
>>> 1 <= characters_q1 <= 5
True
```


Tests for q2_2 (points: [0])
```
>>> 1 <= characters_q2 <= 3
True
```


Tests for q3_1 (points: [0])
```
>>> 1 <= names_q1 <= 4
True
```


Tests for q3_2 (points: [0])
```
>>> 1 <= names_q2 <= 4
True
```


Tests for q3_3 (points: [0])
```
>>> 1 <= names_q3 <= 4
True
```


Tests for q4_1 (points: [0])
```
>>> isinstance(biggest_change, (int, float))
True
```


Tests for q4_2 (points: [0, 0])
```
>>> isinstance(smallest_change_major, (int, float))
True
```
```
>>> 1 <= smallest_change_major <= 3
True
```


Tests for q4_3 (points: [0, 0, 0])
```
>>> isinstance(gws_relative_change, (int, float))
True
```
```
>>> isinstance(linguistics_relative_change, (int, float))
True
```
```
>>> isinstance(rhetoric_relative_change, (int, float))
True
```


Tests for q4_4 (points: [0, 0])
```
>>> isinstance(biggest_rel_change_major, (int, float))
True
```
```
>>> 1 <= biggest_rel_change_major <= 3
True
```


Tests for q6_1 (points: [0])
```
>>> 1 <= survivor_answer <= 3
True
```


Tests for q7_1 (points: [4])
```
>>> contact == 3
True
```


Tests for q7_2 (points: [4])
```
>>> grades == 2
True
```


Tests for q7_3 (points: [4])
```
>>> regrade == 2
True
```


Tests for q7_4 (points: [4])
```
>>> gradescope == 576157
True
```


Tests for q7_5 (points: [4])
```
>>> acceptable == 2
True
```


Tests for q7_6 (points: [4])
```
>>> drops == 2
True
```


Tests for q7_7 (points: [3])
```
>>> exams == 2
True
```


Tests for q7_8 (points: [0, 0])
```
>>> type(secret) == str
True
```
```
>>> len(secret)
8
```


Tests for q8_1 (points: [1])
```
>>> survey == 'data8 is data gr8'
True
```
